# Assignment 3 - Naive Bayes  
## The Iris flowers Dataset


<h3><u/>Classes</u>- types of Iris flower</h3><br/><br/>
<table align='left'>
    <tr><div style="text-align:left;font-size:125%"><u>There are 3 possible flower types:</u></div></tr>
    <tr><th>class 0 - Iris Setosa</th><th>class 1 - Iris Virginica</th><th>class 2 - Iris Versicolor</th></tr>
    <tr>
        <td><img src="./images/Iris-versicolor.jpg" alt="Iris Versicolor" style="width: 200px;text-align:left"/></td>
        <td><img src="./images/iris_setosa.jpg" alt="Iris Setosa" style="width: 200px;text-align:left"/></td>
        <td><img src="./images/iris_virginica.jpg" alt="Iris Virginica" style="width: 200px;text-align:left"/></td>
    </tr>
</table> 

<h3><u>Feature Set</u></h3>
<table align='left' width="200">
    <tr><td>
        <img src="./images/iris_petal_sepal.png" alt="Iris" width="200" align='left'/>
    </td></tr>
    <tr><td align='left'>
        <div style="text-align:left;font-size:125%"><u>Dataset includes 4 features</u>:</div><br/>
        <ul style="list-style-type:circle;text-align:left;font-size:115%">
            <li>sepal-length</li>
            <li>sepal-width</li>
            <li>petal-length</li>
            <li>petal-width</li>
        </ul>
    </td></tr>
</table>    

```python
# --------------------- 
## DO NOT DELETE THE FOLLOWING IMPORTS:
# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------
import pandas as pd
import numpy as np
import sys
%matplotlib inline
# -------
## Add additional import ONLY if needed:
# --------------------- 
```

## General instructions:
Please Read the following instructions, and make sure you follow all steps, before submiting your assignment via moodle:
- Put your code one line after the '# YOUR CODE HERE' remark
- You must remove the 'raise NotImplementedError()' exception raising line (which appears a line after the above remark).
    * if you do not remove this line, it will be a sign that you didn't implement that code, and we won't be able to check your work
- Do NOT remove any other line in this notebook    
- you also need to implement the two functions 'myName', 'myId'
    * myName - you need to return your full name as a string
    * myId - you need to return your ID number as a an integer    
- When you want to check your work, select the 'Cell' --> 'Run All' menu
    * before performing your filnal test, we suggest to clear previous output (by selecting 'Cell' --> 'All output' --> 'Clear' menu) and then reperform the final execution ('Cell' --> 'Run All') of the code.
    * after performing 'Run All', make sure there are no exceptions thrown and that the output is as you expected.
- Don't forget to save your work, by clicking the 'save' icon (in the upper left part of the screen), or by selecting the 'File' --> 'Save and checkpoint' menu item. 
- Do NOT change the file name of this notebook
- The work is done individualy, for each student
- Each student needs to submit his/her assignment through moodle.
- Submit the python notebook only (not any additional possible files) 
<br/><br/>
Good Luck :-)<br/>
The courses staff

### Part 1 - Student information methods:
The following 2 cells consist of 2 methods which aim to validate your details.<br/>
Please <b>make sure to implement them</b>

<!-- grade_id: myNameMethod, points: 0, grade: false, solution: true -->
```python
# --------------------- 
## DO NOT DELETE THIS FUNCTION IT IS NEEDED FOR VALIDATION
# --------------------- 
# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------
# Change the following function to return your full name (as a string).
# example: assume your name is John Smith:
# return 'John Smith'
# --------------------- 
def myName():
    return("Ran Ben-melech")
```

<!-- grade_id: myIdMethod, points: 0, grade: false, solution: true -->
```python
# --------------------- 
## DO NOT DELETE THIS FUNCTION IT IS NEEDED FOR VALIDATION
# --------------------- 
# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------
# Change the following function to return your ID number (as an integer number).
# example: assume your ID number is 1234:
# return 1234 
# --------------------- 
def myId():
    return (204296792)
```

```python
# --------------------------------------------------------
## 'myName' and 'myId' implementation validation
# The Following tests the correct implementation of the 'myName' and 'myId' methods
# This test MUST NOT raise an exception
# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------
aName = myName()
aId = myId()
assert aName is not None, 'no return student name'
assert type(aName) is str, "name is not a string: %r" % aName
print ("'myName' method validation is successfull!")
print ("your name is: %s" %(aName))
print ("--------------------------------------")
assert aId is not None, 'no return student id'
assert type(aId) is int, "id is not an integer: %r" % aId
print ("'myId' method validation is successfull!")
print ("your id is: %d" %(aId))
## END OF 'myName' and 'myId' implementation validation
# --------------------------------------------------------
```

Output:
```
'myName' method validation is successfull!
your name is: Ran Ben-melech
--------------------------------------
'myId' method validation is successfull!
your id is: 204296792
```

### Part 2 - Loading datasets.
The following cells perform 2 things:
* load the Iris dataset
* split the dataset to a train-set and a test-set
##### Note: you NEED to run these cells (but NOT implement anything)

```python
# --------------------------------------------------------
### The Following section loads the iris dataset directly from sklean
# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------
from sklearn import datasets
# --------------------------------------------------------
data_iris = datasets.load_iris()
X_iris = pd.DataFrame(data_iris['data'], columns=data_iris['feature_names'])
y_iris = pd.DataFrame(data_iris['target'], columns=['target'])
# --------------------------------------------------------
# display the iris class information:
# --------------------------------------------------------
classNames = data_iris['target_names']
iris_classes = np.unique(y_iris.values)
strClasses = ["%d (representing '%s')" %(iris_class,classNames[iris_class]) for iris_class in iris_classes]
print ('There are %d Iris classes' %(len(strClasses)))
print ('iris classes:')
print(strClasses)
# --------------------------------------------------------
# display the feature information:
# --------------------------------------------------------
featureNames = [col for col in X_iris.columns]
print ('\nThere are %d features in the feature set (each feature vector has a value for every feature)' %(len(featureNames)))
print('Feature names:')
print(featureNames)
# --------------------------------------------------------
# display the first few rows of the dataset vectors:
# --------------------------------------------------------
print('\nThe dataset includes %d instances (aka feature vectors)' %(X_iris.shape[0]))
print('first few instances:')
print(X_iris.head())
print('\nfirst few corresponding categories:')
print(y_iris.head())
```

Output:
```
There are 3 Iris classes
iris classes:
["0 (representing 'setosa')", "1 (representing 'versicolor')", "2 (representing 'virginica')"]

There are 4 features in the feature set (each feature vector has a value for every feature)
Feature names:
['sepal length (cm)', 'sepal width (cm)', 'petal length (cm)', 'petal width (cm)']

The dataset includes 150 instances (aka feature vectors)
first few instances:
   sepal length (cm)  sepal width (cm)  petal length (cm)  petal width (cm)
0                5.1               3.5                1.4               0.2
1                4.9               3.0                1.4               0.2
2                4.7               3.2                1.3               0.2
3                4.6               3.1                1.5               0.2
4                5.0               3.6                1.4               0.2

first few corresponding categories:
   target
0       0
1       0
2       0
3       0
4       0
```

```python
# --------------------------------------------------------
### The Following section splits the iris dataset directly from sklean
# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------
import sklearn
from sklearn.model_selection import train_test_split

X_train, X_test, y_train, y_test = train_test_split(X_iris, y_iris, test_size=0.2, random_state=42)

# --------------------------------------------------------
# display train-test split information
# --------------------------------------------------------
print ('Information after train-test split:')
print('The train-set includes %d instances and %d corresponding categories\n' %(X_train.shape[0],y_train.shape[0]))
print('The test-set includes %d instances and %d corresponding categories\n' %(X_test.shape[0],y_test.shape[0]))

# --------------------------------------------------------
## concatinate the X_train and y_train for Naive Bayes training:
# --------------------------------------------------------
train_set = pd.concat((X_train, y_train), axis=1)
# --------------------------------------------------------
if 'datasets' in sys.modules:
    del (datasets)
if 'train_test_split' in sys.modules:
    del (train_test_split)
sys_modules = list(sys.modules.keys())
for mdl in sys_modules:
    if mdl.startswith('sklearn.'):
        del(sys.modules[mdl]) 
del (sklearn)
if 'sklearn' in sys.modules:
    del (sys.modules['sklearn'])
# --------------------------------------------------------
# Display the first few rows of the training-set:
# --------------------------------------------------------
print('First few rows of unified train-set:')
train_set.head()
```

Output:
```
Information after train-test split:
The train-set includes 120 instances and 120 corresponding categories

The test-set includes 30 instances and 30 corresponding categories

First few rows of unified train-set:
    sepal length (cm)  sepal width (cm)  petal length (cm)  petal width (cm)  \
22                4.6               3.6                1.0               0.2   
15                5.7               4.4                1.5               0.4   
65                6.7               3.1                4.4               1.4   
11                4.8               3.4                1.6               0.2   
42                4.4               3.2                1.3               0.2   

    target  
22       0  
15       0  
65       1  
11       0  
42       0
```

## The naive bayes Classifier

<img src="./images/bayes.PNG" alt="Naive Bayes Classifier" align='left'/>

## Training a (Guassian) Naive bayes model 
We perform the following during the training step:
1. Calculate Priors
2. Calculate Gaussian Likelihood 'mean' parameter
3. Calculate Gaussian Likelihood 'std' parameter 
4. organize the call to the training steps in the 'fit' method

### Assignment part 3 - train step 1 - calculate category priors:
for each class (0,1,2) you need to calculate the prior.<br/><br/>
<b>prior(y=0)</b>=p(y=0)=count(y=0 in train-set)/count(number-of-instances in train-set)<br/><br/>
<b> do this for each class (0,1,2) </b>

<!-- grade_id: calc-category-priors, points: 0, grade: false, solution: true -->
```python
# --------------------- 
# Change the following function to return a (regular python) list 
#                  containing the 3 classes prirors (one for each category) 
# ------
# input parameter - trainingSet - the train-set dataframe
# ------
# The retuned value is a (regular python) list, containing the priors of the classes
# --------------------- 
def calcCategoryPriors(trainingSet):
    
    yTrain = trainingSet[['target']]
    total=yTrain.shape[0]
    uniqueClasses = np.unique(yTrain['target'].values)
    helpLi=[]
    li=[]
    for m in uniqueClasses:
        helpLi.append(0)
        li.append(0)
    
    for i in yTrain['target']:
        for x in uniqueClasses:
            if i==uniqueClasses[x]:
                helpLi[x]+=1
    for x in range(len(li)):
        li[x]=helpLi[x]/total
    return li
        
```

<!-- grade_id: calc-category-priors-test, points: 2, grade: true, solution: false -->
```python
# --------------------- 
## The Following notebook cell tests the 'calcCategoryPriors' method
# --------------------- 
arrPriors = calcCategoryPriors(train_set)
priorClass_0 = arrPriors[0]
print ('testing for expected class prior ...')
assert int(priorClass_0*100)==33,'wrong prior for class-0'
print ("... 'calcCategoryPriors' test passed successfully :-)")
print ('prior for category 0: %f' %(priorClass_0))
print ("\nIMPORTANT NOTE: additional tests for the 'calcCategoryPriors' method might be performed")
# --------------------- 

```

Output:
```
testing for expected class prior ...
... 'calcCategoryPriors' test passed successfully :-)
prior for category 0: 0.333333

IMPORTANT NOTE: additional tests for the 'calcCategoryPriors' method might be performed
```

### Assignment part 4 - train step 2 - Calculate Gaussian Likelihood 'mean' parameter:
for each feature calculate mean value for the feature in each for each of the class values (0,1,2) seperatly.<br/><br/>
<b>for class 0 (y=0)</b> take the rows consisting 'target' value of 0, and calculate the mean <br/>
To calculate mean use: dataframe[colName].mean() <br/><br/>
<b>Do this for each feature for each class</b>

<!-- grade_id: calc-mean-Gaussian-likelihood, points: 0, grade: false, solution: true -->
```python
# --------------------- 
# Change the following function to return mean of the values of each feature
# this should be done seperatly for each class value (in the training set)
    # create a meanDataframe, including the features as colums and classes as indexes
    # the cells consists of the mean of distribution of feature values (given a specific class)
# ------
# input parameter - trainingSet - the train-set dataframe
# ------
# retuned value - a dataframe, consisting of the feature names as colums and classes as indexes
# the cells consists of the mean of distribution of feature values (given a specific class)
# --------------------- 
def calcMeanLikelihood(trainingSet):
    yTrain = trainingSet[['target']]
    meandf=pd.DataFrame(index=np.unique(yTrain['target'].values),columns=trainingSet.columns)
    meandf=meandf.drop(columns=['target'])
    for i in range(len(meandf)):
        meandf0=trainingSet.loc[trainingSet['target'].values == i]
        meandf0=meandf0.drop(columns=['target'])
        for col in meandf0.columns:
            meandf.loc[i,col]=meandf0[col].mean()
    return meandf
```

<!-- grade_id: calc-mean-test, points: 3, grade: true, solution: false -->
```python
# --------------------- 
## The Following notebook cell tests the 'calcMeanLikelihood' method
# --------------------- 
meanLiklihoodDf = calcMeanLikelihood(train_set)
likelihood_petalLength_class1 = meanLiklihoodDf.iloc[1,2]
print ('testing for expected mean likelihood estimation ...')
assert int(likelihood_petalLength_class1*10)==42,'wrong mean likelihood estimation for petalLength_class1'
print ("... 'calcMeanLikelihood' test passed successfully :-)")
print ('likelihood for the mean of petal length for category 1 is estimated as: %f' %(likelihood_petalLength_class1))
print ("\nIMPORTANT NOTE: additional tests for the 'calcMeanLikelihood' method might be performed")
# --------------------- 

```

Output:
```
testing for expected mean likelihood estimation ...
... 'calcMeanLikelihood' test passed successfully :-)
likelihood for the mean of petal length for category 1 is estimated as: 4.241463

IMPORTANT NOTE: additional tests for the 'calcMeanLikelihood' method might be performed
```

### Assignment part 4 - train step 3 - Calculate Gaussian Likelihood 'std' parameter:
for each feature calculate std value for the feature in each for each of the class values (0,1,2) seperatly.<br/><br/>
<b>for class 0 (y=0)</b> take the rows consisting 'target' value of 0, and calculate the std <br/>
To calculate std use: dataframe[colName].std() <br/><br/>
<b>Do this for each feature for each class</b>

<!-- grade_id: calc-std-Gaussian-likelihood, points: 0, grade: false, solution: true -->
```python
# --------------------- 
# Change the following function to return std (standard diviation) of the values of each feature
# this should be done seperatly for each class value (in the training set)
    # create a stdDataframe, including the features as colums and classes as indexes
    # the cells consists of the std of distribution of feature values (given a specific class)
# ------
# input parameter - trainingSet - the train-set dataframe
# ------
# retuned value - a dataframe, consisting of the feature names as colums and classes as indexes
# the cells consists of the std of distribution of feature values (given a specific class)
# --------------------- 
def calcStdLikelihood(trainingSet):
    yTrain = trainingSet[['target']]
    std=pd.DataFrame(index=np.unique(yTrain['target'].values),columns=trainingSet.columns)
    std=std.drop(columns=['target'])
    for i in range(len(std)):
        std0=trainingSet.loc[trainingSet['target'].values == i]
        std0=std0.drop(columns=['target'])
        for col in std0.columns:
            std.loc[i,col]=std0[col].std()
    return std
```

<!-- grade_id: calc-std-test, points: 3, grade: true, solution: false -->
```python
# --------------------- 
## The Following notebook cell tests the 'calcStdLikelihood' method
# --------------------- 
stdLiklihoodDf = calcStdLikelihood(train_set)
likelihood_petalLength_class1 = stdLiklihoodDf.iloc[1,2]
print ('testing for expected std likelihood estimation ...')
assert int(likelihood_petalLength_class1*10)==4,'wrong std likelihood estimation for petalLength_class1'
print ("... 'calcStdLikelihood' test passed successfully :-)")
print ('likelihood for the std of petal length for category 1 is estimated as: %f' %(likelihood_petalLength_class1))
print ("\nIMPORTANT NOTE: additional tests for the 'calcStdLikelihood' method might be performed")
# --------------------- 

```

Output:
```
testing for expected std likelihood estimation ...
... 'calcStdLikelihood' test passed successfully :-)
likelihood for the std of petal length for category 1 is estimated as: 0.481132

IMPORTANT NOTE: additional tests for the 'calcStdLikelihood' method might be performed
```

### Assignment part 4 - train step 4 - the fit method:
The fit method uses the previous 3 methods for a full (Gaussian) Naive Bayes model training step.<br/>
##### Note: you NEED to run this cell (but NOT implement anything)

```python

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------
def fit(trainingSet):
    """
    1. Calculate the class priors of the training set, using the 'calcCategoryPriors' method.
    2. Calculate the mean of the training set per feature per class, using the 'calcMeanLikelihood' method.
    3. Calculate the std of the training set per feature per class, using the 'stdLiklihoodDf' method.
    """
    arrPriors = calcCategoryPriors(trainingSet)
    meanLiklihoodDf = calcMeanLikelihood(trainingSet)
    stdLiklihoodDf = calcStdLikelihood(trainingSet)
    
    return meanLiklihoodDf, stdLiklihoodDf, arrPriors
```

## Predicting a class for a new example using the (Guassian) Naive bayes model 
We perform the following during the training step:
1. Calculate Guassian likelihood probability, for a given feature value, mean and std
2. Calculate a posteriori probabilities for each training example
3. prdict class for for each training example, given a posteriori probabilities
4. a full predict method using the above

<img src="./images/bayes.PNG" alt="Naive Bayes Classifier" align='left'/>

### Assignment part 5 - predict step 1 - the 'calcGaussianProb' method:
The 'calcGaussianProb' method uses the training methods and returns the Gaussian probablilty <br/>
of that feature value (for a specifc class).<br/>
##### Note: you NEED to run this cell (but NOT implement anything)

<img src="./images/gausianProb.PNG" alt="Gausian likelihood probability" align='left'/>

```python
# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------
"""
given a specific feature value and the trained mean & std (per a specific class) 
We assume normal (Guassian distribution) and we return the density value 
or the Gaussian Probability for that given value
Note: the input parameters are all numbers (scalars)
"""
def calcGaussianProb(xFeatureVal, mean, std):
    exponent = np.exp(-((xFeatureVal-mean)**2 / (2 * std**2 )))
    return (1 / ((2 * np.pi)**(1/2) * std)) * exponent
```

### Assignment part 6 - predict step 2 - the 'calcAposteriorProbs' method:
The 'calcAposteriorProbs' method uses the training parameters to predict the a posteriori probability <br/>
for every test instance, per class <br/>
##### Note: you NEED to run this cell (but NOT implement anything)

```python
# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------
"""
    1. Create a probability matrix to store the results
    2. Update each label's probability using the Gaussian probability function
"""
def calcAposteriorProbs(XTest, arrTrainedClassPriors, dfTrainedMean, dfTrainedStd, categories):
    numClasses = len(categories)
    dfProbPerTestInstPerClass = pd.DataFrame(np.zeros((XTest.shape[0], numClasses)), columns=categories, index=XTest.index)
    for category in categories:
        classPrior = arrTrainedClassPriors[category]
        dfProbPerTestInstPerClass[category]=classPrior
        # Check for each row
        for nRow in range(XTest.shape[0]):

            # Multiply the current given probability by the newly calculated probability for the given event (feature)
            for nCol in range(XTest.shape[1]):
                xFeatureVal=XTest.iloc[nRow, nCol]
                mean=dfTrainedMean.iloc[category,nCol]
                std=dfTrainedStd.iloc[category,nCol]
                gaussianProb = calcGaussianProb(xFeatureVal, mean, std)
                # multiple the prior class probability with the gausian likelihood:
                dfProbPerTestInstPerClass.iloc[nRow, category] *= gaussianProb
    return dfProbPerTestInstPerClass            
```

### Assignment part 7 - predict step 3 - the 'predictClasses' method:
The 'predictClasses' method uses the calculated  a posteriori probabilites <br/>
for every test instance, to calculate the most probable class for each test instance <br/>


<!-- grade_id: predict-classes, points: 0, grade: false, solution: true -->
```python
# --------------------- 
# Change the following function to return a Series data structure, consisting of 
#                  a predicted class for each test instance, using the calculated a posterior probabilities
# ------
# input parameter - df_probPerTestInstPerClass - a dataframe of the a posterior probabilities, for each test instance,
#                                               note 1: columns are the classes (for each instance the probability
#                                                        for class 0 is in column 0, for class 1 in column 1, etc.)
#                                               note 2: the indexes correspond to the xTest indexes
# ------
# The retuned value is a Series data-structure, with the same index, and with the predicted class in each cell
# --------------------- 
def predictClasses(df_probPerTestInstPerClass):
    res=pd.Series(index=df_probPerTestInstPerClass.index)
    for row in df_probPerTestInstPerClass.index:
        res[row]=df_probPerTestInstPerClass.loc[row].idxmax()
    print (res)
    return res
    
```

### Assignment part 8 - predict step 4 - the 'predict' method:
The 'predict' method is a Guasian Naive Bayes classifier <br/>
It uses the above methods to predict test instances <br/>


```python
# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------
# --------------------- 
# Change the following function to return the predicted categories of the test set examples
# one category for each correcponding test instance
# this should be done seperatly for each class value (in the training set)
    # create a stdDataframe, including the features as colums and classes as indexes
    # the cells consists of the std of distribution of feature values (given a specific class)
# ------
# input parameter - XTest - the test-set instances
# input parameter - arrTrainedClassPriors - the trained priors of the classes
# input parameter - dfTrainedMean - the dataframe of the trained means for each feature per each class
# input parameter - dfTrainedStd - the dataframe of the trained stds for each feature per each class
# input parameter - categories - the possible categories (in the iris case includes [0,1,2])
# ------
# retuned value - a dataframe, consisting of the feature names as colums and classes as indexes
# the cells consists of the std of distribution of feature values (given a specific class)
# --------------------- 
def predict(XTest, arrTrainedClassPriors, dfTrainedMean, dfTrainedStd, categories):
    # 1. calculate a posterior probabities:
    dfProbPerTestInstPerClass = calcAposteriorProbs(XTest, arrTrainedClassPriors, dfTrainedMean, dfTrainedStd, categories)

    # 2. predict classes using the a posterior probabities:
    results = predictClasses(dfProbPerTestInstPerClass)
    
    return results
```

```python
def evaluate_accuracy(y_true, y_pred):
    """
    Compare how many predictions were correct (compare the y_hat to y)
    """
    accuracy_score = pd.Series(y_true.values == y_pred.values).value_counts() * 100 / y_true.shape[0]
    return accuracy_score.iloc[0]
```

<!-- grade_id: test-precition, points: 2, grade: true, solution: false -->
```python
# --------------------- 
## The Following tests the the predict, using the accuracy function
# --------------------- 
meanLiklihoodDf, stdLiklihoodDf, arrPriors=fit(train_set)
iris_classes = np.unique(train_set['target'].values)
y_hat = predict(X_test, arrPriors, meanLiklihoodDf, stdLiklihoodDf, iris_classes)
accuracy_score = evaluate_accuracy(y_test['target'], y_hat)
assert accuracy_score == 100, "accuracy should be 100"
print("Accuracy Score: {}".format(accuracy_score))
```

Output:
```
73     1.0
18     0.0
118    2.0
78     1.0
76     1.0
31     0.0
64     1.0
141    2.0
68     1.0
82     1.0
110    2.0
12     0.0
36     0.0
9      0.0
19     0.0
56     1.0
104    2.0
69     1.0
55     1.0
132    2.0
29     0.0
127    2.0
26     0.0
128    2.0
131    2.0
145    2.0
108    2.0
143    2.0
45     0.0
30     0.0
dtype: float64
Accuracy Score: 100.0
```